---
title       : Operações vetoriais
description : Explorando propriedades dos vetores
--- type:NormalExercise lang:r xp:100 skills:1 key:b37009c969
## Reciclagem


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp:100 skills:1 key:0a4c024393
## Coerção 


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
